```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

<a id='verytop'></a>

# Homework 1 - v2
## Prerequisite Review & Intro to Jupyter Notebooks and Python
## Due Date: Saturday, Jan 20th, 11:59 PM on Gradescope


### Detailed Submission Instructions Are Provided at the end of this Notebook


## Assignment Scope:


1. Navigate Jupyter notebooks (like this one);
2. Write and evaluate some basic *expressions* in Python, the computer language of the course;
3. Practice with basic  *functions*
4. Practice with the NumPy module
5. Review prerequisite topics from Calculus and Discrete Structures
6. Complete the course Welcome Survey



## Collaboration Policy

Data science is a collaborative activity.  However a key step in learning and retention is **creating solutions on your own.**  

Below are examples of acceptable vs unacceptable use of resources and collaboration when doing HW assignments in CSCI 3022.


The following would be some **examples of cheating** when working on HW assignments in CSCI 3022.  Any of these constitute a **violation of the course's collaboration policy and will result in an F in the course and a trip to the honor council**.   


 - Consulting web pages that may have a solution to a given homework problem or one similar is cheating.  However, consulting the class notes, and web pages that explain the material taught in class but do NOT show a solution to the homework problem in question are permissible to view.  Clearly, there's a fuzzy line here between a valid use of resources and cheating. To avoid this line, one should merely consult the course notes, the course textbook, and references that contain syntax and/or formulas.
 - Copying a segment of code or math solution of three lines or more from another student from a printout, handwritten copy, or by looking at their computer screen 
 - Allowing another student to copy a segment of your code or math solution of three lines or more
 - Taking a copy of another student's work (or a solution found online) and then editing that copy
 - Reading someone else’s solution to a problem on the HW before writing your own.
 - Asking someone to write all or part of a program or solution for you.
 - Asking someone else for the code necessary to fix the error for you, other than for simple syntactical errors
 


On the other hand, the following are some **examples of things which would NOT usually be
considered to be cheating**:
 - Working on a HW problem on your own first and then discussing with a classmate a particular part in the problem solution where you are stuck.  After clarifying any questions you should then continue to write your solution independently.
 - Asking someone (or searching online) how a particular construct in the language works.
 - Asking someone (or searching online) how to formulate a particular construct in the language.
 - Asking someone for help in finding an error in your program.  
 - Asking someone why a particular construct does not work as you expected in a given program.
   

To test whether you are truly doing your own work and retaining what you've learned you should be able to easily reproduce from scratch and explain a HW solution that was your own when asked in office hours by a TA/Instructor or on a quiz/exam.   


If you have difficulty in formulating the general solution to a problem on your own, or
you have difficulty in translating that general solution into a program, it is advisable to see
your instructor or teaching assistant rather than another student as this situation can easily
lead to a, possibly inadvertent, cheating situation.

We are here to help!  Visit HW Hours and/or post questions on Piazza!


## Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to either in-notebook tests and/or hidden tests that won't be run until you upload to Gradescope.  

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown

Question | Points | Grading Type
--- | --- | ---
Question 1 | 10 | manual
Question 2 | 2 | autograded
Question 3 | 2 | autograded
Question 4 | 2 | autograded
Question 5 | 4 | autograded
Question 6 | 3 | autograded
Question 7 | 2 | autograded
Question 8 | 2 | autograded
Question 9 | 4 | autograded & manually graded
Question 10 | 3 | manual
Question 11 | 11 | manual
Question 12 |5 | autograded
<a id='top'></a>Total | 50

## Recommended Readings for this HW:

 - [Python for Data Analysis: Python Language Basics](https://wesmckinney.com/book/python-basics#tut_basics)
 - [Python for Data Analysis: Dictionaries](https://wesmckinney.com/book/python-builtin#dict)
 - [Computation and Inferential Thinking:  Intro to Chapter 8: Functions](https://inferentialthinking.com/chapters/08/Functions_and_Tables.html) 
 - [Python Like You Mean It:  Vectorization Using NumPy](https://www.pythonlikeyoumeanit.com/Module3_IntroducingNumpy/VectorizedOperations.html)
 - [Prerequisite Review Resources posted on the  Modules in Canvas](https://canvas.colorado.edu/courses/101142/pages/prerequisite-review-calculus-and-discrete-structures?module_item_id=5144193) 


<a id='top'></a>
---
**Shortcuts:**   [Problem 2](#p2) | [Problem 3](#p3) | [Problem 4](#p4) | [Problem 5](#p5) | [Problem 6](#p6) | [Problem 7](#p7) | [Problem 8](#p8)| [Problem 9](#p9) | [Problem 10](#p10) | [Problem 11](#p11) | [Problem 12](#p12) |

---




## <span style='color:Red'>  Question 1 (10 pts):  
Complete the [Intro to CSCI 3022 Video Assignment available on Canvas](https://canvas.colorado.edu/courses/101142/assignments/1926004)  (to receive points you must watch the entire video and answer the embedded questions).

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


# 1. Jupyter Notebooks
This webpage is called a Jupyter Notebook.  Jupyter Notebook is an open source web application that you can use to create and share documents that contain live code, equations, visualizations, and text.  
(Fun fact: The name, Jupyter, is a loose acronym meaning Julia, Python, and R, although now it supports many other languages).     
We’ll be using Python in this class. 

Throughout the semester you can access all Juptyer notebooks for this class via the CU JupyterHub (i.e. you do NOT need to download any software locally on your computer)



### Jupyter Shortcuts ###

Here are some useful Jupyter notebook keyboard shortcuts.  To learn more keyboard shortcuts, go to **Help -> Keyboard Shortcuts** in the menu above. 

Here are a few we like:
1. `ctrl`+`return` : *Evaluate the current cell*
1. `shift`+`return`: *Evaluate the current cell and move to the next*
1. `esc` : *command mode* (may need to press before using any of the commands below)
1. `esc`+`a` : *create a cell above*
1. `esc`+`b` : *create a cell below*
1. `esc`+`dd` : *delete a cell*
1. `esc`+`m` : *convert a cell to markdown*
1. `esc`+`y` : *convert a cell to code*

## 1.1. Text cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings. 

After you edit a text cell, click the "run cell" button at the top that looks like ▶| or hold down `shift` + `return` to confirm any changes. 

**Practice 1.1.1.** This paragraph is in its own text cell.  Try editing it so that this sentence is the last sentence in the paragraph, and then click the "run cell" ▶| button or hold down `shift` + `return`.  This sentence, for example, should be deleted.  So should this one.

**Practice 1.1.2**:   Double click on this cell to see the Markdown "behind the scenes"</span>:
 
This is type

<span style='color:Green'> This is GREEN type </span>

<span style='background :yellow' > This is highlighted type </span>

### This is big type
## This is bigger type
# This is even bigger type


*Italic*  or _Italic_ 

**Bold** 

***Bold and Italic*** 

~~Scratch Me.~~


## 1.2. Code cells
Other cells contain code in the Python 3 language. Running a code cell will execute all of the code it contains.

To run the code in a code cell, first click on that cell to activate it.  It'll be highlighted with a little green or blue rectangle.  Next, either press ▶| or hold down `shift` + `return`.

**Practice 1.2.1** Try running the cell below:

```python
print("Hello, World!")
```

Output:
```
Hello, World!
```

And this one:

```python
print("\N{WAVING HAND SIGN}, \N{EARTH GLOBE ASIA-AUSTRALIA}!")
```

Output:
```
👋, 🌏!
```

**Expressions:**  The fundamental building block of Python code is an expression. Cells can contain multiple lines with multiple expressions. When you run a cell, the lines of code are executed in the order in which they appear. Every `print` expression prints a line. 

Run the next cell and notice the order of the output.

```python
print("First this line is printed,")
print("and then this one.")
```

Output:
```
First this line is printed,
and then this one.
```

## 1.3. Writing Jupyter notebooks
You can use Jupyter notebooks for your own projects or documents.  When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, click the + button in the menu bar.  It'll start out as a text cell.  You can change it to a code cell by clicking inside it so it's highlighted, clicking the drop-down box next to the restart (⟳) button in the menu bar, and choosing "Code".

**Practice 1.3.1.** Add a code cell below this one.  Write code in it that prints out:
   

    First this line,
    then the whole 🌏,
    and then this one.


Run your cell to verify that it works.

```python
print("First this line,")
print("then the whole \N{EARTH GLOBE ASIA-AUSTRALIA},")
print("and then this one.")
```

Output:
```
First this line,
then the whole 🌏,
and then this one.
```

## 1.4. Errors
Python is a language, and like natural human languages, it has rules.  It differs from natural language in two important ways:
1. The rules are *simple*.  You can learn most of them in a few weeks and gain reasonable proficiency with the language in a semester.
2. The rules are *rigid*.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes.  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong.

Errors are okay; even experienced programmers make many errors.  When you make an error, you just have to find the source of the problem, fix it, and move on.

```python
print("This line is missing something.")
```

Output:
```
This line is missing something.
```

**Note:** In the toolbar, there is the option to click `Cell > Run All`, which will run all the code cells in this notebook in order. However, the notebook stops running code cells if it hits an error, like the one in the cell above.

You should see something like this (minus our annotations):

<img src="img/error.jpg"/>

The last line of the error output attempts to tell you what went wrong.  The *syntax* of a language is its structure, and this `SyntaxError` tells you that you have created an illegal structure.  The message is saying Python expected you to write something more (in this case, a right parenthesis) before finishing the cell.

There's a lot of terminology in programming languages, but you don't need to know it all in order to program effectively. If you see a cryptic message like this, you can often get by without deciphering it.  (Of course, if you're frustrated, ask a neighbor or a staff member for help.) 


**Practice 1.4.1** Fix the code above so that you can run the cell and see the intended message instead of an error.

## 1.5. The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

Next to every code cell, you'll see some text that says `[ ]`. Before you run the cell, you'll see `[ ]`. When the cell is running, you'll see `[*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `[1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection.  Or if you find that when you try to run the top cell in this notebook it outputs an error it's an issue with the kernel.  If this happens, try the following steps:

1. At the top of your screen, click **Kernel**, then **Interrupt**.
2. If that doesn't help, click **Kernel**, then **Restart**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.

# 2. Real Number Types in Python

Quantitative information arises everywhere in data science. In addition to representing commands to print out lines, expressions can represent numbers and methods of combining numbers.


Python has two real number types:

`int`: 		an integer of any size

`float`:	a number with an optional fractional part

An int never has a decimal point; a float always does

A float might be printed using scientific notation




The expression `3.2500` evaluates to the floating point 3.25. (Run the cell and see.)

```python
3.2500
```

Output:
```
3.25
```

Notice that we didn't have to `print`. When you run a notebook cell, if the last line has a value, then Jupyter helpfully prints out that value for you. However, it won't print out prior lines automatically.

```python
print(2)
3
4
```

Output:
```
2
4
```

Above, you should see that 4 is the value of the last expression, 2 is printed, but 3 is lost forever because it was neither printed nor last.

You don't want to print everything all the time anyway.  But if you feel sorry for 3, change the cell above to print it.

## 2.1 Limitations of float values:

Be aware that float values have the following limitations, which can lead to issues if you're not aware:

- They have limited precision of 15-16 decimal places

- After arithmetic, the final few decimal places can be wrong

- They have limited size (but the limit is huge)


**Practice 2.2.1**:  Run the cells below to observe some of the issues that arise when working with float values

```python
382764238764.238742034876503865238723985 - 382764238764.238742
```

Output:
```
0.0
```

```python
print(0.1 + 0.2)
```

Output:
```
0.30000000000000004
```

```python
print(0.1+.1+.1)
```

Output:
```
False
```

# 3. Names and Assignment Statements in Python
In natural language, we have terminology that lets us quickly reference very complicated concepts.  We don't say, "That's a large mammal with brown fur and sharp teeth!"  Instead, we just say, "Bear!"

In Python, we do this with *assignment statements*. An assignment statement has a name on the left side of an `=` sign and an expression to be evaluated on the right.

```python
ten = 3 * 2 + 4
```

When you run that cell, Python first computes the value of the expression on the right-hand side, `3 * 2 + 4`, which is the number 10.  Then it assigns that value to the name `ten`.  At that point, the code in the cell is done running.

After you run that cell, the value 10 is bound to the name `ten`:

```python
ten
```

Output:
```
10
```

The statement `ten = 3 * 2 + 4` is not asserting that `ten` is already equal to `3 * 2 + 4`, as we might expect by analogy with math notation.  Rather, that line of code changes what `ten` means; it now refers to the value 10, whereas before it meant nothing at all.

If the designers of Python had been ruthlessly pedantic, they might have made us write

    define the name ten to hereafter have the value of 3 * 2 + 4 

instead.  You will probably appreciate the brevity of "`=`"!  But keep in mind that this is the real meaning.

**Practice 3.1.** Run the following cell which uses a variable name `eleven` that hasn't been assigned to anything. You'll see an error!

```python
eleven = 0
eleven + 8
```

Output:
```
8
```

A common pattern in Jupyter notebooks is to assign a value to a name and then immediately evaluate the name in the last line in the cell so that the value is displayed as output. 

```python
close_to_pi = 355/113
close_to_pi
```

Output:
```
3.1415929203539825
```

Another common pattern is that a series of lines in a single cell will build up a complex computation in stages, naming the intermediate results.

```python
semimonthly_salary = 842.5
monthly_salary = 2 * semimonthly_salary
number_of_months_in_a_year = 12
yearly_salary = number_of_months_in_a_year * monthly_salary
yearly_salary
```

Output:
```
20220.0
```

Names in Python can have letters (upper- and lower-case letters are both okay and count as different letters), underscores, and numbers.  The first character can't be a number (otherwise a name might look like a number).  And names can't contain spaces, since spaces are used to separate pieces of code from each other.

Other than those rules, what you name something doesn't matter *to Python*.  For example, this cell does the same thing as the above cell, except everything has a different name:

```python
a = 842.5
b = 2 * a
c = 12
d = c * b
d
```

Output:
```
20220.0
```

**However**, names are very important for making your code *readable* to yourself and others.  The cell above is shorter, but it's totally useless without an explanation of what it does.

## 3.1. Checking Your Code

Now that you know how to name things, you can start using the built-in *tests* to check whether your work is correct. Sometimes, there are multiple tests for a single question, and passing all of them is required to receive credit for the question. 


If you haven't already, scroll up to the top of this notebook and run the first cell that says `Initialize Otter:`   [Back to top](#verytop) 


Then run the cell below.  It will import the module NumPy which will be used in some of the tests below.  We will discuss more about NumPy and importing modules in a later section of this assignment.

```python
# Imports module NumPy - we will explore this module in a later section of this assignment.
import numpy as np

# Function needed to run in-notebook tests
import hashlib

def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()

```

<a id='p2'></a>

## <span style='color:Red'>  Question 2 (2 points)

Assign the name `seconds_in_a_decade` to the number of seconds between midnight January 1, 2010 and midnight January 1, 2020. Note that there are two leap years in this span of a decade. A non-leap year has 365 days and a leap year has 366 days.

Replace the ellipses (`...`) with your expression and run the cell.

Running the cell directly after it will test whether you have assigned `seconds_in_a_decade` to the correct answer.


```python
# Change the next line 
# so that it computes the number of seconds in a decade 
# and assigns that number the name, seconds_in_a_decade.

seconds_in_a_decade = (8*365*24*60*60) + (2*366*24*60*60)
# We've put this line in this cell 
# so that it will print the value you've given to seconds_in_a_decade when you run it.  
# You don't need to change this.
seconds_in_a_decade
```

Output:
```
315532800
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

[Back to top](#top)

<a id='p3'></a>


## <span style='color:Red'>   Question 3 (2 pts)
When you run the following cell, Python produces a cryptic error message.

```python
two = 1
```

**Multiple Choice:**
Choose the best explanation of what's wrong with the code, and then assign the the variable  `names_q4` to the corresponding letter (as a capital letter string `'A'`,`'B'`, `'C'`,or `'D'`) to  below to indicate your answer.  Replace the ellipses (`...`) with your answer.  

A. Python is smart and already knows `4 = 2 + 2`.

B. It should be `2 + 2 = 4`.

C. In Python, it's a rule that the `=` sign must have a variable name to its left, and `4` isn't a variable name.

D. I don't get an error message. This is a trick question.


```python
names_q3 = 'C'

# We've put this next line in this cell 
# so that it will print the value you've given to names_q3 when you run it.  
# You don't need to change this:
names_q3

#Note that the grader check below only 
#checks if you have entered a valid input (i.e. either a capital 'A', 'B', 'C', or 'D')
#not whether or not you chose the correct input - 
#that will be graded via a hidden test once you submit to Gradescope.
```

Output:
```
'C'
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## 3.2. Operators


Many basic arithmetic operations are built into Python.  The textbook section on [Expressions](http://www.inferentialthinking.com/chapters/03/1/Expressions.html) describes all some key arithmetic operators.  The common operator that differs from typical math notation is `**`, which raises one number to the power of the other. So, `2**3` stands for $2^3$ and evaluates to 8. 

**Practice:** Play around with the operators $\%, /, **$, and $//$ to see what they do:

```python

```

The order of operations is the same as what you learned in elementary school, and Python also has parentheses.  For example, compare the outputs of the cells below. The second cell uses parentheses for a happy new school year!

```python
7+6*5-6*3**2*2**3/4*7
```

Output:
```
-719.0
```

```python
7+(6*5-(6*3))**2*((2**3)/4*7)
```

Output:
```
2023.0
```

In standard math notation, the first expression is

$$6 + 6 \times 5 - 6 \times 3^2 \times \frac{2^3}{4} \times 7,$$

while the second expression is

$$6 + (6 \times 5 - (6 \times 3))^2 \times \left(\frac{\left(2^3\right)}{4} \times 7\right).$$


[Back to top](#top)

<a id='p4'></a>
    
## <span style='color:Red'>  Question 4 (2 pts) ###

Write a Python expression in this next cell that's equal to $\displaystyle 5 \times \left(3 \frac{10}{11}\right) - 50 \frac{1}{3} + 2^{0.5 \times 22} - \frac{7}{33} + 7$.  That's five times three and ten elevenths, minus fifty and a third, plus two to the power of half twenty-two, minus seven thirty-thirds plus seven.  By "$\displaystyle 3 \frac{10}{11}$" we mean $\displaystyle 3+\frac{10}{11}$, not $\displaystyle 3 \times \frac{10}{11}$.

Replace the ellipses (`...`) with your expression.  Try to use parentheses only when necessary.



```python
q4_answer = 5*(3+(10/11))-(50+(1/3))+(2**(.5*22))-(7/33)+7

# We've put this line in this cell 
# so that it will print the value you've given to q4_answer when you run it.  
# You don't need to change this.
q4_answer
```

Output:
```
2024.0
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## 3.3: Useful Python Data Structure: Lists

A list is a type of Python object that allows us to store a sequence of other objects. One of its major utilities is that it provides us with means for updating the contents of a list later on.

A list object is created using square-brackets, and its contents are separated by commas: [item1, item2, ..., itemN]. 

**Its contents need not be of the same type of object.**

```python
# we can store whatever we want in lists!
my_list = [2,1,5,'squid']
```

```python
# Python indexing starts at 0
my_list[0]
```

Output:
```
2
```

In order to access a range of elements in a list, you need to slice a list. One way to do this is to use the simple slicing operator i.e. colon(:). With this operator, one can specify where to start the slicing, where to end, and specify the step.  

For list slicing in Python the ending index is EXCLUSIVE but the beginning index is INCLUSIVE

```python
# and array/list slicing in Python ending index is EXCLUSIVE,
# but the beginning index is INCLUSIVE
my_list[1:3]
```

Output:
```
[1, 5]
```

```python
# you can even reference elements from the end of the list!
my_list[-1]
```

Output:
```
'squid'
```

```python
my_list[-2]
```

Output:
```
5
```

```python
my_list[-2:]
```

Output:
```
[5, 'squid']
```

## 3.5:  Useful Python Data Structure:  Dictionaries

Python’s dictionary allows you to store key-value pairs, and then pass the dictionary a key to quickly retrieve its corresponding value. Specifically, you construct the dictionary by specifying one-way mappings from key-objects to value-objects. 

Each key must map to exactly one value, meaning that a key must be unique.

```python
# Dictionaries take the following key-value form
my_dict = {
    'key1': 'value1',
    'key2': 'value2'
}
```

```python
# we can assign a new key-value pair to our dictionary like so
my_dict['key3'] = 'porcupine'
```

```python
# dictionaries can store any kind of object as its value, including lists, tuples, or even other dictionaries
my_list = [3.5, 2, "anything"]

my_dict['my_list'] = my_list

my_dict
```

Output:
```
{'key1': 'value1',
 'key2': 'value2',
 'key3': 'porcupine',
 'my_list': [3.5, 2, 'anything']}
```

```python
# we can access a dictionaries contents by using the stored key value
my_dict['key3']
```

Output:
```
'porcupine'
```

```python
# we can also view the contents of our dictionary by calling the keys() and values() methods
print(my_dict.keys())
print(my_dict.values())
```

Output:
```
dict_keys(['key1', 'key2', 'key3', 'my_list'])
dict_values(['value1', 'value2', 'porcupine', [3.5, 2, 'anything']])
```

```python
# we can overwrite a value in the dictionary in the same way that we assigned a value
my_dict['key1'] = 'new value!'
my_dict['key1']
```

Output:
```
'new value!'
```

It is important to note outright that the time it takes for dictionary to take a key and retrieve a value does not depend on the size of the dictionary. That is the complexity for a dictionary look-up is 
! It accomplishes this by making use of a technique known as hashing (https://en.wikipedia.org/wiki/Hash_function)



## 3.5 Loops

You can loop over indices like you would in C++ or Fortran (if you are suitably tough) by using `range`. Note that `range(3)` makes use of the fact that Python is 0-based and will loop over something of length 3: [0, 1, 2]

```python
for i in range(3):
    print(i)
```

Output:
```
0
1
2
```

```python
print(list(range(3)))
```

Output:
```
[0, 1, 2]
```

You can also loop over the elements of a list:

```python
x = [2, 7, -153, 'octopus']

for xi in x:
    print(xi)
```

Output:
```
2
7
-153
octopus
```

```python
# or just loop over the indices
for i in range(len(x)):
    print(x[i])
```

Output:
```
2
7
-153
octopus
```

```python
list(range(1,5))
```

Output:
```
[1, 2, 3, 4]
```

## 3.6 List Comprehensions

Rather than using loops to create a list, we will frequently use a list comprehension, as they take less time to run:

List comprehensions have the form: 
`[f(x) for x in seq]`




**Practice:**  Suppose you have a list of fruits and you want a new list, containing only the fruits with the letter "a" in the name.

```python
fruits = ["apple", "banana", "cherry", "kiwi", "mango"]
```

Without list comprehension you will have to write a for statement with a conditional test inside:

```python
newlist = []

for x in fruits:
  if "a" in x:
    newlist.append(x)

newlist
```

Output:
```
['apple', 'banana', 'mango']
```

With list comprehension you can do all that with only one line of code:

```python
newlist = [x for x in fruits if "a" in x]

newlist
```

Output:
```
['apple', 'banana', 'mango']
```

## 3.7 Comments
You may have noticed these lines in the cell in which you answered Practice 1.3.1:

    # Change the next line 
    # so that it computes the number of seconds in a decade 
    # and assigns that number the name, seconds_in_a_decade.
    
This is called a *comment*. It doesn't make anything happen in Python; Python ignores anything on a line after a `#`.  Instead, it's there to communicate something about the code to you, the human reader. Comments are extremely useful. 

# 4. Functions

The most common way to combine or manipulate values in Python is by calling functions. 


Functions are either built-in, or user-defined as demonstrated below. 


## 4.1 Built In Functions:
Python comes with many built-in functions that perform common operations.

Examples of a few built-in functions we will frequently use include:  
`abs(), sum(), min(), max(), len(), set(), type(), map(), int() lambda(), round()`


For example, the `abs` function takes a single number as its argument and returns the absolute value of that number. Run the next two cells and see if you understand the output.

```python
abs(5)
```

Output:
```
5
```

```python
abs(-5)
```

Output:
```
5
```

You can check the type of any object in Python using the built-in `type` function:

```python
type(2)
```

Output:
```
int
```

```python
type(2.5)
```

Output:
```
float
```

```python
type('hello')
```

Output:
```
str
```

```python
type([1,2.5,'hello'])
```

Output:
```
list
```

## 4.2 Multiple arguments
Some functions take multiple arguments, separated by commas. For example, the built-in `max` function returns the maximum argument passed to it.

```python
max(2, -3, 4, -5)
```

Output:
```
4
```

## 4.3 Viewing Documentation

To output the documentation for a function, use the `help` function.

```python
help(max)
```

Output:
```
Help on built-in function max in module builtins:

max(...)
    max(iterable, *[, default=obj, key=func]) -> value
    max(arg1, arg2, *args, *[, key=func]) -> value
    
    With a single iterable argument, return its biggest item. The
    default keyword-only argument specifies an object to return if
    the provided iterable is empty.
    With two or more arguments, return the largest argument.
```

#### You can also use Jupyter to view function documentation inside your notebook: 
(The function must already be defined in the kernel for this to work).

**Practice 4.1.1:** Below, click your mouse anywhere on the `sum` block below and use `Shift` + `Tab` to view the function's documentation. 

```python
max(2, -3, 4, -5)
```

Output:
```
4
```

## 4.4. User Designed Functions


Let's write a very simple function that converts a proportion to a percentage by multiplying it by 100.  

For example, the value of `to_percentage(.5)` should be the number 50 (no percent sign).

A function definition has a few parts.

##### `def`
It always starts with `def` (short for **def**ine):

    def

##### Name
Next comes the name of the function.  Like other names we've defined, it can't start with a number or contain spaces. Let's call our function `to_percentage`:
    
    def to_percentage

##### Signature
Next comes something called the *signature* of the function.  This tells Python how many arguments your function should have, and what names you'll use to refer to those arguments in the function's code.  A function can have any number of arguments (including 0!). 

`to_percentage` should take one argument, and we'll call that argument `proportion` since it should be a proportion.

    def to_percentage(proportion)
    
If we want our function to take more than one argument, we add a comma between each argument name. Note that if we had zero arguments, we'd still place the parentheses () after that name. 

We put a **colon** after the signature to tell Python that the next indented lines are the body of the function. If you're getting a syntax error after defining a function, check to make sure you remembered the colon!

    def to_percentage(proportion):

##### Documentation
Functions can do complicated things, so you should write an explanation of what your function does.  For small functions, this is less important, but it's a good habit to learn from the start.  Conventionally, Python functions are documented by writing an **indented** triple-quoted string:

    def to_percentage(proportion):
        """Converts a proportion to a percentage."""

##### Body
Now we start writing code that runs when the function is called.  This is called the *body* of the function and every line **must be indented with a tab**.  Any lines that are *not* indented and left-aligned with the def statement is considered outside the function. 

Some notes about the body of the function:
- We can write code that we would write anywhere else.  
- We use the arguments defined in the function signature. We can do this because we assume that when we call the function, values are already assigned to those arguments.
- We generally avoid referencing variables defined *outside* the function. If you would like to reference variables outside of the function, pass them through as arguments!


Now, let's give a name to the number we multiply a proportion by to get a percentage:

    def to_percentage(proportion):
        """Converts a proportion to a percentage."""
        factor = 100

##### `return`
The special instruction `return` is part of the function's body and tells Python to make the value of the function call equal to whatever comes right after `return`.  We want the value of `to_percentage(.5)` to be the proportion .5 times the factor 100, so we write:

    def to_percentage(proportion):
        """Converts a proportion to a percentage."""
        factor = 100
        return proportion * factor
        
`return` only makes sense in the context of a function, and **can never be used outside of a function**. `return` is always the last line of the function because Python stops executing the body of a function once it hits a `return` statement. If a function does not have a return statement, it will not return anything; if you expect a value back from the function, make sure to include a return statement. 

*Note:*  `return` inside a function tells Python what value the function evaluates to. However, there are other functions, like `print`, that have no `return` value. For example, `print` simply prints a certain value out to the console. 

In short, `return` is used when you want to tell the *computer* what the value of some variable is, while `print` is used to tell you, a *human*, its value.

[Back to top](#top)

<a id='p5'></a>

## <span style='color:Red'>   Question 5a (2 pts) ###

Define `to_percentage` in the cell below. 

Call your function to convert the proportion 0.2 to a percentage.  Name that percentage `twenty_percent`.

Replace EACH ellipses below (`...`)  with the corresponding correct code.  


```python
def to_percentage(x):
    return (x*100)

twenty_percent = to_percentage(.2)
twenty_percent
```

Output:
```
20.0
```

```python
grader.check("q5a")
```

Output:
```
q5a results: All test cases passed!
```

## 4.5 Variable Scope
Here's something important about functions: the names assigned *within* a function body are only accessible within the function body. Once the function has returned, those names are gone.  So even if you created a variable called `factor` and defined `factor = 100` inside of the body of the `to_percentage` function and then called `to_percentage`, `factor` would not have a value assigned to it outside of the body of `to_percentage`:

**Note**: Below, you should see an error message.

```python
# You should get an error when you run this.  (If you don't, 
# you might have defined factor somewhere above.)
factor = 1
```


## <span style='color:Red'>  Question 5b (2 pts):

You can also pass in named values as arguments to your function.

Use `to_percentage` again to convert the proportion named `a_proportion` (defined below) to a percentage called `a_percentage`.

Your answer should call the function `to_percentage`

*Note:* You don't need to define `to_percentage` again!  Like other named values, functions stick around after you define them.

```python
a_proportion = 2**(0.5) / 2
a_percentage = to_percentage(a_proportion)
a_percentage
```

Output:
```
70.71067811865476
```

```python
grader.check("q5b")
```

Output:
```
q5b results: All test cases passed!
```

[Back to top](#top)

<a id='p6'></a>

## <span style='color:Red'>  Question 6 (3 pts):

Write a function `square` that takes a number as its argument and returns the square of its argument.

```python
def square(x):
    return (x*x)

```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

## 5. Functions vs. Methods vs. Attributes

An important characteristic of the Python language is the consistency of its object model. Every number, string, data structure, function, class, module, and so on exists in the Python interpreter in its own “box,” which is referred to as a Python object. Each object has an associated type (e.g., integer, string, or function) and internal data. In practice this makes the language very flexible, as even functions can be treated like any other object.

There are some distinctions between Functions, Methods and Attributes:


#### Functions

Functions are either built-in, or user-defined as demonstrated above. Functions are called directly (e.g., type(x), help(str), etc)


```python
x='hello!'

type(x)
```

Output:
```
str
```

```python
square(6)
```

Output:
```
36
```

```python
help(square) # notice our doc string is printed when we call the help function
```

Output:
```
Help on function square in module __main__:

square(x)
```

```python
type(square)
```

Output:
```
function
```

## 5.1 Methods


A method is a function associated with a specified object (and, by extension, with the type of that object). That is, a method corresponds to a data-type operation.  Methods are defined within the class from which an object is instantiated. 


We can call (or invoke) a method as follows:  `obj.method_name()`


**Practice 5.1.1** As a simple example, Python's built-in int type has a method named bit_length(), so you can determine the number of bits in the binary representation of an int value as follows:


```python
x = 3 ** 100
bits = x.bit_length()
bits
```

Output:
```
159
```

```python
my_list
```

Output:
```
[3.5, 2, 'anything']
```

```python
my_list.append(1)
my_list
```

Output:
```
[3.5, 2, 'anything', 1]
```

```python
my_dict.keys()
```

Output:
```
dict_keys(['key1', 'key2', 'key3', 'my_list'])
```

## 5.2 Attributes

Attributes are variables that belong to an object. Attributes are defined within the class from which an object is instantiated. Attributes hold information about the object, and can be accessed by referencing the object and then calling the attribute by name. Note that many attributes for built-in Python classes are "dundered", meaning that the name is prefixed and suffixed with a double underscore (e.g. `plus_some.__name__`

Attributes are very useful when designing your own Python classes and when working with other libraries such as Numpy and Pandas


You can use the `dir` function to see all the attributes and methods associated with a given object

```python
dir(my_list)
```

Output:
```
['__add__',
 '__class__',
 '__class_getitem__',
 '__contains__',
 '__delattr__',
 '__delitem__',
 '__dir__',
 '__doc__',
 '__eq__',
 '__format__',
 '__ge__',
 '__getattribute__',
 '__getitem__',
 '__gt__',
 '__hash__',
 '__iadd__',
 '__imul__',
 '__init__',
 '__init_subclass__',
 '__iter__',
 '__le__',
 '__len__',
 '__lt__',
 '__mul__',
 '__ne__',
 '__new__',
 '__reduce__',
 '__reduce_ex__',
 '__repr__',
 '__reversed__',
 '__rmul__',
 '__setattr__',
 '__setitem__',
 '__sizeof__',
 '__str__',
 '__subclasshook__',
 'append',
 'clear',
 'copy',
 'count',
 'extend',
 'index',
 'insert',
 'pop',
 'remove',
 'reverse',
 'sort']
```

```python
#recall
my_list
```

Output:
```
[3.5, 2, 'anything', 1]
```

```python
#example of a list attribute
my_list.__len__()
```

Output:
```
4
```

# 6. Importing Modules

Most programming involves work that is very similar to work that has been done before.  Since writing code is time-consuming, it's good to rely on others' published code when you can.  Rather than copy-pasting, Python allows us to **import modules**. A module is a file with Python code that has defined variables and functions. By importing a module, we are able to use its code in our own notebook.

Python includes many useful modules that are just an `import` away.  


## 6.1:  NumPy module:

𝐍𝐮𝐦𝐏𝐲  (short for "Numeric Python") is an open source Python module for scientific computing. NumPy supports large, multidimensional arrays and matrices. It also supports a large collection of mathematical functions not found in Python's standard math library.

It's custom in data science to import Numpy  with the alias $\texttt{np}$.  We can then ***access any mathematical function in the Numpy library by prepending function names by $\texttt{np}$***

For a more in-depth description of methods and attributes built into Numpy see here:  https://numpy.org/doc/stable/user/quickstart.html   
and  here:  https://numpy.org/doc/stable/reference/routines.math.html

*Hint: If you press `shift+tab` while next to the function call, the documentation for that function will appear.*

By using the `as` keyword, you can give imports different variable names.

In Sec 3.1 above, the line `import numpy as np` imports the numpy module. 

This statement creates a module and then assigns the name `np` to that module. We are now able to access any variables or functions defined within `np` by typing the name of the module followed by a dot, then followed by the name of the variable or function we want.

    <np>.<name>

```python
radius = 5
area_of_circle = radius**2 * np.pi
area_of_circle
```

Output:
```
78.53981633974483
```

[Back to top](#top)

<a id='p7'></a>


## <span style='color:Red'>  Question 7 (2 pts)
The module `numpy` also provides the name `e` for the base of the natural logarithm, which is roughly 2.71. 

Compute $e^{\pi}-\pi$, giving it the name `near_twenty`.

*Remember: You can access `pi` from the `numpy` module as well!*


```python
near_twenty = (np.e**np.pi) - np.pi
near_twenty
```

Output:
```
19.99909997918947
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

## 6.2  Accessing Functions

In the question above, you accessed variables within the `numpy` module. 

**Modules** also define **functions**.  For example, `numpy` provides the name `floor` for the floor function.  Having imported `numpy as np` already, we can write `np.floor(7.5)` to compute the floor of 7.5.  (Note that the floor function returns the largest integer less than or equal to a given number.)


```python
np.floor(7.5)
```

Output:
```
7.0
```

[Back to top](#top)

<a id='p8'></a>

## <span style='color:Red'> Question 8 (2 pts)

Calculate the square root of `pi` using `sqrt` and `pi` from the `numpy` module.  Give the result the name `sqrt_of_pi`.

```python
sqrt_of_pi = np.sqrt(np.pi)
sqrt_of_pi
```

Output:
```
1.7724538509055159
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

## 6.3 NumPy Arrays

The core of NumPy is the array. Like Python lists, arrays store data; however, they store data in a more efficient manner. In many cases, this allows for faster computation and data manipulation.

Use `np.array` to create an array. It takes a sequence, such as a list or range. 

Below, we create an array `arr` containing the values 1, 2, 3, 4, and 5 (in that order).

```python
arr = np.array([1,2,3, 4, 5])
```

In addition to values in the array, we can access attributes such as shape and data type. A full list of attributes can be found [here](https://docs.scipy.org/doc/numpy-1.15.0/reference/arrays.ndarray.html#array-attributes).

```python
arr[3]
```

Output:
```
4
```

```python
arr[2:4]
```

Output:
```
array([3, 4])
```

```python
arr.shape
```

Output:
```
(5,)
```

```python
arr.dtype
```

Output:
```
dtype('int64')
```

Arrays, unlike Python lists, cannot store items of different data types.

```python
# A regular Python list can store items of different data types
[1, '3']
```

Output:
```
[1, '3']
```

```python
# Arrays will convert everything to the same data type
np.array([1, '3'])
```

Output:
```
array(['1', '3'], dtype='<U21')
```

```python
# Another example of array type conversion
np.array([5, 8.3])
```

Output:
```
array([5. , 8.3])
```

### 6.3.1 Making Arrays of Zeros

```python
np.zeros(5)
```

Output:
```
array([0., 0., 0., 0., 0.])
```

### 6.3.2. Making Arrays from ranges:

The `np.arange(start, stop, step)` function creates an array that starts at the indicated number.  It is exclusive of the stopping number (i.e. it won't include the stopping number).  And it uses the stepsize indicated.

```python
np.arange(0, 10, 2)
```

Output:
```
array([0, 2, 4, 6, 8])
```

## 6.3.3 Interpolating numbers 

The `linspace(start,end,num)` function generates `num` numbers evenly spaced between the `start` and `end` (again INCLUSIVE of start and EXCLUSIVE of end)

```python
np.linspace(0, 5, 11)
```

Output:
```
array([0. , 0.5, 1. , 1.5, 2. , 2.5, 3. , 3.5, 4. , 4.5, 5. ])
```

## 6.4 Vectorized Operations in NumPy 

NumPy arrays can be utilized in mathematical expressions to perform mathematical computations using an array’s entries. In general, NumPy implements mathematical functions such that, when a function acts on an array, the mathematical operation is applied to each entry in the array.


For example, observe the following:

```python
# multiply each number in an array by 11
11 * np.array([2, 3, 4])
```

Output:
```
array([22, 33, 44])
```

```python
# divide each number in an array by 50
np.array([25, 40, 50])/50
```

Output:
```
array([0.5, 0.8, 1. ])
```

```python
# NumPy array addition will add them element-wise
np.array([1, 2, 3]) + np.array([4, 5, 6])
```

Output:
```
array([5, 7, 9])
```

```python
# Whereas Python list addition will concatenate the two lists
[1, 2, 3] + [4, 5, 6]
```

Output:
```
[1, 2, 3, 4, 5, 6]
```

```python
# Here's how you can take the square root of each element of an array WITHOUT using a for-loop:

arr = np.array([1, 4, 9, 16, 25])

np.sqrt(arr)
```

Output:
```
array([1., 2., 3., 4., 5.])
```


### Vectorized Operations:
NumPy’s arrays are homogeneous: an array can only contain data of a single type. For instance, an array can contain 8-bit integers or 32-bit floating point numbers, but not a mix of the two. This is in stark contrast to Python’s lists and tuples, which are entirely unrestricted in the variety of contents they can possess; a given list could simultaneously contain strings, integers, and other objects. 

This restriction on an array’s contents comes at a great benefit; in “knowing” that an array’s contents are homogeneous in data type, NumPy is able to delegate the task of performing mathematical operations on the array’s contents to optimized, compiled C code. This is a process that is referred to as **vectorization**. 

The outcome of this can be a tremendous speedup relative to the analogous computation performed in Python, which must painstakingly check the data type of every one of the items as it iterates over the arrays, since Python typically works with lists with unrestricted contents.

**Definition**:  In the context of a high-level languages like Python, the **term vectorization** describes the use of optimized, pre-compiled code written in a low-level language (e.g. C) to **perform mathematical operations over a sequence of data. This is done in place of an explicit iteration written in the native language code (e.g. a “for-loop” written in Python).**


Consider, for instance, the task of summing the integers 0-9,999 stored in an array. Calling NumPy’s `sum` function cues optimized C code to iterate over the integers in the array and tally the sum. `np.sum` is therefore a “vectorized” function. 


Let’s compare how long it takes to compute this sum using a loop via vectorization using a [Jupyter magic command](https://ipython.readthedocs.io/en/stable/interactive/magics.html) called `%%time`, which times the execution of that cell. 

You can use this by writing it as the first line of a cell. (Note that `%%` is used for *cell magic commands* that apply to the entire cell, whereas `%` is used for *line magic commands* that only apply to a single line.)


```python
%%time

# sum an array, using NumPy's vectorized 'sum' function
np.sum(np.arange(10000))  
49995000
```

Output:
```
CPU times: user 128 µs, sys: 47 µs, total: 175 µs
Wall time: 148 µs
49995000
```

Now let’s compare this to the time required to explicitly loop over the array in Python and tally up the sum. Python is unable to take advantage of the fact that the array’s contents are all of a single data type - it has to check, for every iteration, if it is dealing with an integer, a string, a floating point number, etc, just as it does when iterating over a list. This will slow down the computation massively.

```python
%%time

total = 0
for i in np.arange(10000):
    total = i + total

    total
49995000
```

Output:
```
CPU times: user 1.67 ms, sys: 0 ns, total: 1.67 ms
Wall time: 1.47 ms
49995000
```

The sum is significantly faster when performed in using NumPy’s vectorized function! This should make it clear that, whenever computational efficiency is important, one should avoid performing explicit for-loops over long sequences of data in Python, be them lists or NumPy arrays. NumPy provides a whole suite of vectorized functions. In fact, the name of the game when it comes to leveraging NumPy to do computations over arrays of numbers is to exclusively leverage its vectorized functions. 


For more details on vectorization using NumPy see this reading:
 https://www.pythonlikeyoumeanit.com/Module3_IntroducingNumpy/VectorizedOperations.html)

<!-- BEGIN QUESTION -->

[Back to top](#top)

<a id='p9'></a>


## <span style='color:Red'>  Question 9 (5 pts)

Write a function `summation(n)` that uses vectorization in Numpy to evaluate the following summation for $n \geq 1$:

$$\sum_{i=1}^{n} \left( i^3 + 3 i^2 \right)$$

**Note**: You should **NOT** use ANY `for` loops in your solution. You may find `np.arange` helpful for this question!

```python
def summation(n):
    return np.sum((np.power(np.arange(1,n+1,1),2))*3) + np.sum((np.power(np.arange(1,n+1,1),3)))
#Do not change the following cells:
print("summation(5) =  ", summation(5))
print("summation(200) = ", summation(200))
    
```

Output:
```
summation(5) =   390
summation(200) =  412070100
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

<!-- END QUESTION -->

## 6.5 Methods vs Functions in NumPy

Many things are implemented in NumPy as both functions and methods:

Examples are `mean`, `min`, `max`, `sum`, `reshape`

For example, there is a `np.sum` function, that adds up all the elements:

```python
arr = np.array([1, 2, 0, 1])
np.sum(arr)
4
```

Output:
```
4
```

There is also a `sum()` method of the numpy array object:

```python
type(arr)
```

Output:
```
numpy.ndarray
```

```python
arr.sum()
```

Output:
```
4
```

In these cases choosing the method or the function will usually depend on which one is easier to read.

# 7. Math Prerequisite Review

### 7.1 How to Answer Math Questions in this Class

$\text{Below is an example question to demonstrate the type of work and justification}$
$\text{that is required for full points when answering math questions using LaTex}$.
## $\color{blue}{\text{EXAMPLE QUESTION:}}  $

Write a full solution to the following question using LaTex (i.e. do NOT use code to calculate the answer).

What values of $x$ solve the equation: $\phantom{xx}x^2-5x=-6$?

### $\color{blue}{\text{This correct answer would be worth zero points}}:$

$x=2$ and $x=3$

### $\color{blue}{\text{This answer would be worth full points}}:$


Solution to Example Question:

$x^2-5x=-6$ 

$\implies x^2-5x+6=0$

$\implies (x-2)(x-3)=0$

$\implies \boxed{x=2}$ or $\boxed{x=3}$.

### 7.2 LaTeX 
You should use LaTeX to format math in your answers. If you aren't familiar with LaTeX, not to worry ([ here are some tips.](https://ds100.org/sp22/resources/assets/hw/latex_tips.pdf)
) 

It's not hard to use in a Jupyter notebook. Just place your math in between dollar signs **WITHOUT ANY SPACES before or after the dollar signs**:

\\$f(x) = 2x\\$ becomes $f(x) = 2x$.

If you have a longer equation, use double dollar signs to place it on a line by itself:

\\$\\$\sum_{i=0}^n i^2\\$\\$ becomes:

$$\sum_{i=0}^n i^2$$.

Here is some handy LaTex:

| Output | Latex   |
|:--|:--|
| $$x^{a + b}$$  | `x^{a + b}` |
| $$x_{a + b}$$ | `x_{a + b}` |
| $$\frac{a}{b}$$ | `\frac{a}{b}` |
| $$\sqrt{a + b}$$ | `\sqrt{a + b}` |
| $$\{ \alpha, \beta, \gamma, \pi, \mu, \sigma^2  \}$$ | `\{ \alpha, \beta, \gamma, \pi, \mu, \sigma^2  \}` |
| $$\int_{x=2}^{\infty} \frac{1}{x} \,dx $$ | `\sum_{x=2}^{\infty} \frac{1}{x} \,dx` |
| $$\sum_{x=1}^{100} x$$ | `\sum_{x=1}^{100} x` |
| $$\frac{d}{d x} $$ | `\frac{d}{dx} ` |
| $$ \frac{x^3}{3} \biggr|_{a}^{b} $$ | Double click on this cell to see the LaTeX|
| Formatting aligned equations:        \begin{align}
y &= \int_0^4 xe^{xt} \,dt \\
x &= \ln(u) \\
\end{align}  |  You can align multiple lines using the & anchor, \\ newline, in an align block [as shown in this link](https://www.overleaf.com/learn/latex/Aligning_equations_with_amsmath#Aligning_several_equations) Note if you use align DO NOT also add $ notation.


Here is an excellent [LaTeX cheat sheet](https://pbil.univ-lyon1.fr/members/mbailly/Comm_Scientifique/docs/latexsheet.pdf) you can reference. 


[Back to top](#top)

<a id='p10'></a>


## Prerequisite Review: Calculus 

Exponential and logarithmic functions and their integrals and derivatives will be used in this class when we discuss probability distributions.  


See the [Prerequisite Review Resources posted on the  Modules in Canvas](https://canvas.colorado.edu/courses/101142/pages/prerequisite-review-calculus-and-discrete-structures?module_item_id=5144193) if you need to review any of these key concepts.


## <span style='color:Red'> Question 10 (3 pts):

To begin to brush up on these topics, evaluate the following improper integral, assuming $\lambda>0$. Show all steps. 

$$\int_{0}^{\infty} \lambda e^{-\lambda x} \,dx$$

Hint:
Recall to evaluate an improper integral start by replacing the infinite upper limit with a variable and take the limit as the variable goes to infinity (that way you can use the Fundamental Theorem of Calculus when evaluating the integral):

$$\int_{0}^{\infty} \lambda e^{-\lambda x} \,dx$$

$$ = \lim_{b \to \infty} \int_0^b \lambda e^{-\lambda x} \,dx$$

$$ = ... $$


Note that we could have also written the lines of math above using using LaTeX `align`, which allows us to select where to align each row.  Double click on this cell to see the code for this.  Note that you DON'T use $$ notation if you decide to use the align option- the pdf won't compile at the end of this HW if you have both.  Also, there can't be any blank lines in between the align statements in the LaTeX code.  If you decide to us this option see [this link](https://www.overleaf.com/learn/latex/Aligning_equations_with_amsmath#Aligning_several_equations)  for more details about using align 

\begin{align}
&  \int_{0}^{\infty} \lambda e^{-\lambda x} \,dx  \\
& =  \lim_{b \to \infty} \int_0^b \lambda e^{-\lambda x} \,dx \\
& = ... 
\end{align}



<!-- BEGIN QUESTION -->

#### Question 10 Solution) 

**Type your answer to Question 10** in this cell (You can copy the LaTeX from the first 2 lines above as a start and then complete the rest of the problem).  Show all of your steps and fully justify your answer.  Do not add any additional cells to this part. 
\begin{align}
&  \int_{0}^{\infty} \lambda e^{-\lambda x} \,dx  \\
& =  \lim_{b \to \infty} \int_0^b \lambda e^{-\lambda x} \,dx \\
& = \lim_{b \to \infty} \lambda \int_0^b e^{-\lambda x} \,dx \\
& = \lim_{b \to \infty} \lambda \int_0^b e^{u} \,dx \\
& u = {-\lambda x} \rightarrow \frac {du}{dx} = {-\lambda} \rightarrow \frac{du}{-\lambda} = dx\\
& = \lim_{b \to \infty}- \int_0^b e^{u} \,du \\
& = \lim_{b \to \infty} \ - (e^{-\lambda x} \biggr|_{0}^{b})\\
& = \lim_{b \to \infty} \ - (e^{-\lambda x} -e^{0}) \\
& = - (e^{-\infty} -e^{0})\\
& = - (0 - 1)\\
& = 1
\end{align}
This problem was solved through the use of substitution. First I simplified the problem by pulling lambda out of the integrand, and then I began to solve through substitution. After taking the integral I plugged in the u value in and finished the math.

<!-- END QUESTION -->

Q10 Final Answer Check).  To check your final answer to 10, enter the final answer you came up with in the cell below. Note that this is just a built-in public test so you can check your work and determine if you are on the right track.  To receive credit on this problem you must show all steps in part 10 above using LaTeX and fully justifying your answer using correct mathematical notation.  

```python
q10_answer = 1
```

```python
grader.check("q10check")
```

Output:
```
q10check results: All test cases passed!
```

## Prerequisite Review: Conditional Probability 

In Discrete Structures, you were introduced to many key probability concepts.  One that we will frequently use in this course is the concept of Conditional Probability.  

If you need a refresher, see the [Prerequisite Review Resources posted on the  Modules in Canvas](https://canvas.colorado.edu/courses/101142/pages/prerequisite-review-calculus-and-discrete-structures?module_item_id=5144193) to review this concept.  

[Back to top](#top)

<a id='p11'></a>

## <span style='color:Red'>  Question 11 (11 pts):
 
a).  What is the mathematical definition of the conditional probability of an event A given an event B?  Give your answer using LaTex and correct mathematical notation.  (Hint: If you need a refresher, see the links provided on Canvas to review Prerequisites)


b). You flip a fair coin twice.  What is the probability that both flips are heads given that at least one of the flips is a heads?  Explain/justify using math.


c).  Suppose a bag has 5 red marbles and 3 blue ones.  We select 2 marbles without replacement (i.e. we select the first marble and then we don't put it back and then select a 2nd marble). 
    
 - i).  What is the probability that the 2nd marble is red given that the first marble is red?  Explain/justify using math. .
 - ii).  What is the probability that both marbles are blue?  Explain/justify using math.  
    
    


<!-- BEGIN QUESTION -->

#### Question 11abc Solutions) 
Use LaTeX (not code) in the cells below to show all of your steps for parts 11a, 11b and 11c and fully justify your answers.   Do not add any additional cells to this part.

**Enter your answer for part 11a)**
\begin{align}
\frac{p({Event A})}{p({Event B})}
\end{align}
Event A is conditional to event B therefore sits in the numerator.

**Enter your answer for part 11b)**
\begin{align}
\frac{p({2 heads})}{p({1 head})} \rightarrow \frac{1/4}{3/4} = \frac{1}{3}
\end{align}
Place the probabilities relative to one another and perform calculations.

**Enter your answer for part 11ci)** 
\begin{align}
\frac{amountRed}{Total} \rightarrow \frac{4}{7}
\end{align}
Considering that there is now one less marble in red and one less marble overall, we can adjust our numbers accordingly to reveal the probabilirt of random selection here.

**Enter your answer for part 11cii)** 
\begin{align}
p(blueSelection) \cdot p(blueSelection2 | blueSelection) \rightarrow \frac{3}{8} \cdot \frac{2}{7} = \frac{3}{28}
\end{align}
First, we calculate the possibility of of selecting a blue marble. Then, because the second event is dependent on the first one, we caclulate the selection of the marble as if the first event occurred and then multiply the two.

<!-- END QUESTION -->

11b Final Answer Check).  To check your final answers to 11b and 11c, enter the final answers you came up (as fractions) with in the cell below. Note that this is just a built-in public test so you can check your work and determine if you are on the right track.  To receive credit on this problem you must show all steps in part 12b above using LaTeX and fully justifying your answer using correct mathematical notation.  

```python
q11b_answer = 1/3

q11ci_answer = 4/7

q11cii_answer = 3/28

```

```python
grader.check("q11check")
```

Output:
```
q11check results: All test cases passed!
```

[Back to top](#top)

<a id='p12'></a>

## <span style='color:Red'>  Question 12 (5 pts):  Welcome Survey


If you haven't already, please fill out the [CSCI 3022 Welcome Survey at this link:](https://forms.gle/xRzHeWyhzezLa7Tm7)


At the end of the survey, there will be a 5-digit integer code displayed. Assign `q12_survey` to that integer below (i.e. replace the ellipses (`...`) with your 5-digit code). **After you have typed your answer in the cell run the cell (hold down shift + return)**.  


Running the `grader.check` cell directly after it will test whether you have assigned `q12_survey` to the correct integer code. 


```python
q12_survey = 30728
```

```python
grader.check("q12")
```

Output:
```
q12 results: All test cases passed!
```


<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 1!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any websites you referenced*

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 


* **Homework 1 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW1 Autograded assignment in Gradescope.

* **Homework 1 Manually Graded**: Submit your hw01.PDF to the HW1 Manually Graded assignment in Gradescope.  **It is your responsibility to fully review your PDF file before submitting and make sure that all your lines of code are visible and any LaTeX has correctly compiled and is fully viewable.**  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.** If not, you will lose points.    

**TROUBLESHOOTING TIPS**  If you are having any issues compiling your assignment, please [read through these troubleshooting tips first](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing), then post any questions on Piazza.  

**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw01.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Output:
```
Running your submission against local test cases...



Your submission received the following results when run against available test cases:

    q2 results: All test cases passed!

    q3 results: All test cases passed!

    q4 results: All test cases passed!

    q5a results: All test cases passed!

    q5b results: All test cases passed!

    q6 results: All test cases passed!

    q7 results: All test cases passed!

    q8 results: All test cases passed!

    q9 results: All test cases passed!

    q10check results: All test cases passed!

    q11check results: All test cases passed!

    q12 results: All test cases passed!
<IPython.core.display.HTML object>
```

Tests for q10check (points: 0)
```
>>> assert get_hash(q10_answer) == 'c4ca4238a0b923820dcc509a6f75849b'
```


Tests for q11check (points: 0)
```
>>> assert get_hash(q11b_answer) == '75b8c9e60e79d5e9519de89fe0893397'
```
```
>>> assert get_hash(q11ci_answer) == '63f7bd5af06cede8abe8385027a9e238'
```
```
>>> assert get_hash(q11cii_answer)=='137eb4aa09cbba8867ca5b31ce37a9e3'
```


Tests for q12 (points: 5)
```
>>> # TEST
>>> assert get_hash(q12_survey) == '6b52f0c01e26aec069e5bd6701dabf8b'
```


Tests for q2 (points: 2)
```
>>> assert get_hash(seconds_in_a_decade) == '7b54be8a1e43a89b0736f55aec0290c5'
```


Tests for q3 (points: [0])
```
>>> # TEST
>>> 
>>> assert names_q3 in ['A', 'B', 'C', 'D']
```


Tests for q4 (points: 2)
```
>>> assert get_hash(q4_answer) == '48325466001978e098c8d0cb216d7cef'
```


Tests for q5a (points: 2)
```
>>> # Make sure your function has the proper syntax!
>>> assert get_hash(to_percentage(0.35))=='e2e536cebe0aee01493dbece3a42cf40'
```
```
>>> # Make sure you are passing in the proper argument!
>>> assert get_hash(twenty_percent) == '75cf3ac5e70c76583be3efb5012bd44e'
```


Tests for q5b (points: 2)
```
>>> assert get_hash(round(a_percentage,4))=='d5283852114ab721760809b6c091520f'
```


Tests for q6 (points: [0.5, 0.5, 2])
```
>>> assert get_hash(square(1)) == 'c4ca4238a0b923820dcc509a6f75849b'
```
```
>>> assert get_hash(square(0)) == 'cfcd208495d565ef66e7dff9f98764da'
```
```
>>> assert get_hash(square(2.5)) == 'd62c232f0b3e8fb5da7db1ff61509b94'
```


Tests for q7 (points: 2)
```
>>> assert get_hash(round(near_twenty,4))=='bc8dee77249d35f9baa359d1c2ab36b4'
```


Tests for q8 (points: 2)
```
>>> assert get_hash(round(sqrt_of_pi,4)) == '4ccf5d819bcfa0bd1603dc31ab5aa649'
```


Tests for q9 (points: 1)
```
>>> # TEST
>>> assert get_hash(summation(1)) == 'a87ff679a2f3e71d9181a67b7542122c'
```
```
>>> # TEST
>>> assert get_hash(summation(2)) == '1ff1de774005f8da13f42943881c655f'
```
```
>>> # TEST
>>> assert get_hash(summation(100)) == '932348b9cc333ed91b5153c91769d66c'
```
